---
title: "Iteración"
output: 
  learnr::tutorial:
    language: es
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Utiliza iteración para reducir duplicación en tu código.
---

<!-- Add JavaScript code for making the exercise code larger -->
<script language="JavaScript" src="js/exercise-font-size.js"></script>

```{r setup, include=FALSE}
# load packages-----------------------------------------------------------------

library(learnr)
library(tidyverse)
library(datos)
library(gradethis)
knitr::opts_chunk$set(echo = FALSE, exercise.checker = gradethis::grade_learnr)

# set options for exercises and checking ---------------------------------------
pinguino_medidas <- tibble(
  pico = c(100,200,300),
  aletas= c(200,400,600),
  cabeza = c(400,800,1200))

mi_media_sin_nas <- function(x) {
  x %>%
    mean(na.rm = TRUE)}

# hide non-exercise code chunks ------------------------------------------------
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(echo = FALSE)
```

## Introducción

```{r imagen-peng, fig.margin = TRUE, echo = FALSE, out.width = "100%", fig.cap="Image from unsplash: Author Andrea Gradilone"}
knitr::include_graphics("https://images.unsplash.com/photo-1506842605307-bbcbd9ea121f?ixlib=rb-1.2.1&ixid=MnwxMjA3fDB8MHxwaG90by1wYWdlfHx8fGVufDB8fHx8&auto=format&fit=crop&w=1267&q=80")
```

### Objetivos de aprendizaje

En este tutorial vas a aprender como resolver problemas de iteración en R. En concreto, vas a aprender a diferenciar como resolver este tipo de problemas con:

* Bucles
* Funciones preconstruidas para situaciones especificas como `colMeans()` 
* Funciones echas a medida con `purrr`

Al final de este tutorial habrás practicado con varias herramientas de programación funcional en R.

## Repetición manual

Vamos a trabajar con unos nuevos datos de pinguinos llamados `pinguino_medidas` con variables, `pico`, `aletas` y `cabeza`. 

```{r pinguino-medidas, echo=TRUE}
pinguino_medidas
```

Imagina que quieres hacer la media de cada una de las tres variables de `pinguino_medidas`. Calcula la media de cada una de las columnas por separado:

```{r meansep,exercise = TRUE}
___(pinguino_medidas$pico)
___(pinguino_medidas$aletas)
___(pinguino_medidas$cabeza)
```

```{r meansep-hint-1}
mean(pinguino_medidas$pico)
___(pinguino_medidas$aletas)
___(pinguino_medidas$cabeza)
```

```{r meansep-hint-2}
mean(pinguino_medidas$pico)
mean(pinguino_medidas$aletas)
___(pinguino_medidas$cabeza)
```

```{r mean-sep-hint-3}
mean(pinguino_medidas$pico)
mean(pinguino_medidas$aletas)
mean(pinguino_medidas$cabeza)
```

```{r meansep-solution}
mean(pinguino_medidas$pico)
mean(pinguino_medidas$aletas)
mean(pinguino_medidas$cabeza)
```

###

Puedes calcular las medias de esta manera, pero estarías repitiendo tu código (recuerda la norma de oro de no copiar y pegar más de dos veces una misma ejecución). Para evitar esto, podrías usar **un bucle _for_**.

## Iteración con bucles


Recueda las 3 cosas que tiene un bucle _for_:

* `output`: La parte del código que asigna suficiente espacio para guardar el resultado. 
* `secuencia`: La parte del código que determina cuantas veces hay que ejecutar el bucle.
* `cuerpo`: La parte del código que hace el trabajo.

Podrías construir un bucle para iterar sobre los `pinguino_medidas` y calcular la media de cada variable.

```{r, echo=TRUE}
salida_media <- vector("double", ncol(pinguino_medidas)) #output

for (i in seq_along(pinguino_medidas)) {     #sequencia
    salida_media[[i]] <- mean(pinguino_medidas[[i]])}     #cuerpo  
```

```{r, echo=TRUE}
salida_media
```

###

Imagina que en lugar de la media quieres calcular la mediana de cada variable, ¿cómo serías tu bucle _for_? ¡Inténtalo!

```{r ex1, exercise = TRUE}
salida_mediana <- vector("double", ncol(pinguino_medidas)) 

for (i in seq_along(pinguino_medidas)) {    
    salida_media[[i]] <- ___(pinguino_medidas[[i]])}      
salida_mediana
```

```{r ex1-solution}
salida_mediana <- vector("double", ncol(pinguino_medidas)) 

for (i in seq_along(pinguino_medidas)) {   
    salida_mediana[[i]] <- median(pinguino_medidas[[i]])} 
salida_mediana
```

###

A pesar de que los bucles __for__ nos permiten reducir el código, resultan más engorrosos porque hay que inicializar el output al comienzo, especificar el número de veces que el bucle se va a ejecutar, y escribir el bucle en sí. A continuación vamos a ver como podemos usar vectorización para simplificar este trabajo.


## Vectorización

En lugar de escribir un bucle __for__, podrías valerte del poder de [vectorizacion en R](https://swcarpentry.github.io/r-novice-gapminder-es/09-vectorization/) y calcular la media de todas las columnas con la function predefinida `colMeans()`:

```{r, echo=TRUE}
colMeans(pinguino_medidas)
```
Como ves `colMeans()`  ejecuta la media de todas las columnas del dataset, y de esta manera no es necesario escribir un bucle for. Esta función utiliza los principios de vectorización, lo que hace la escritura de código más concisa, fácil de leer y menos propensa a errores.

## Purrr

El paquete `purrr` ofrece otra alternativa más flexible que nos permite ejecutar iteraciones _sin tener que escribir bucles for_. Por ejemplo, puedes calcular la media de todas las columnas de `pinguino_medidas` así:

```{r, echo=TRUE}
map(pinguino_medidas,mean)
```

###

El diagrama de abajo te ayudará a recordar como functiona map. En el ejemplo de `pinguino_medidas.` map aplica la functión `mean()` a todas las variables de pinguino_medidas.

```{r echo=FALSE, out.width = "70%"}
knitr::include_graphics("https://raw.githubusercontent.com/rstudio-education/primers/master/iterate/map/www/images/map.png")
```

También puedes utilizar map con el pipe:

```{r, echo=TRUE}
pinguino_medidas %>% map(mean)
```

###

¿Qué tendrías que poner para calcular la mediana de todas las variables de `pinguino_medidas`?

```{r ex2, exercise = TRUE}
pinguino_medidas %>% map(___)

```

```{r ex2-solution}
pinguino_medidas %>% map(median)
```

```{r ex2-check}
grade_code("En este caso la media y la mediana tienen el mismo valor.")
```

###

También puedes pasar tu propia función a map. Por ejemplo, puedes crear una función que calcule la media después de quitar NAs. 

```{r, echo= TRUE}
mi_media_sin_nas <- function(x) {
  x %>%
    mean(na.rm = TRUE)}
```

Y calcula la media para variables numéricas de pinguinos con tu nueva función.

```{r, echo=TRUE}
pinguinos %>% 
  select_if(is.numeric) %>%
  map(mi_media_sin_nas)
```

### Las funciones map

Si has estado atenta, habrás visto que map retorna por defecto una lista. Hay otras functiones de map que devuelven distintos tipos de output.

Función map    | Output                                               | Ejemplo
----------------- | ------------------------------------------------- | -------
**map**             | Lista                         | `map(pinguino_medidas, mean)`
**map_dbl**             | Vector de dobles                  | `map_dbl(pinguino_medidas, mean)`
**map_chr**             | Vector de caracteres                         | `map_chr(pinguino_medidas, mean)`
**map_int**             | Vector de enteros                         | `map_int(pinguino_medidas, length)`
**map_lgl**             | Vector lógico                         | `map_lgl(pinguino_medidas, is.numeric)`


¿Cómo podrías retornar el resultado anterior como un vector de dobles?

```{r ex3, exercise = TRUE}
pinguinos %>% select_if(is.numeric) %>% ___(mi_media_sin_nas)
```

```{r ex3-solution}
pinguinos %>% select_if(is.numeric) %>% map_dbl(mi_media_sin_nas)
```

## En resumen

Enhorabuena. Has practicado iteración con bucles for y con map().

map() tiene más cosas útiles como atajos de sintaxis. Puedes practicar como utilizar tajos de sintaxis en  [este tutorial en ingles](https://rstudio.cloud/learn/primers/5.3).

