---
title: "Iteración"
output: 
  learnr::tutorial:
    language: es
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Utiliza iteración para reducir duplicación en tu código.
---

<!-- Add JavaScript code for making the exercise code larger -->
<script language="JavaScript" src="js/exercise-font-size.js"></script>

```{r setup, include=FALSE}
# load packages----------------------------------------------------------------
library(learnr)
library(tidyverse)
library(datos)
library(gradethis)
library(purrr)
knitr::opts_chunk$set(echo = FALSE, exercise.checker = gradethis::grade_learnr)

# set options for exercises and checking ---------------------------------------
pinguino_medidas <- tibble(
  pico = c(100,200,300),
  aletas= c(200,400,600),
  cabeza = c(400,800,1200))

pinguino_lista_Torgersen <- list(
  masa_corporal_g = c(4433, 4338, 4559),
  anio= c(2017,2018,2019),
  cabeza = c(450,850,1250))

pinguino_lista_Biscoe <- list(
  masa_corporal_g = c(4868, 3504, 5219),
  anio= c(2017,2018,2019),
  cabeza = c(903, 1125,  514))

pinguino_lista_Dream <- list(
  masa_corporal_g = c(4280, 5134, 5183),
  anio= c(2017,2018,2019),
  cabeza = c(588,  718, 1094))

pinguino_lista <- list(
  pinguino_lista_Torgersen,
  pinguino_lista_Biscoe,
  pinguino_lista_Dream
)

names(pinguino_lista)<-c("Torgersen","Biscoe","Dream")


mi_media_sin_nas <- function(x) {
  x %>%
    mean(na.rm = TRUE)}

# hide non-exercise code chunks ------------------------------------------------
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(echo = FALSE)
```

## Introducción

```{r imagen-peng, fig.margin = TRUE, echo = FALSE, out.width = "100%", fig.cap="Image from unsplash: Author Andrea Gradilone"}
knitr::include_graphics("https://images.unsplash.com/photo-1506842605307-bbcbd9ea121f?ixlib=rb-1.2.1&ixid=MnwxMjA3fDB8MHxwaG90by1wYWdlfHx8fGVufDB8fHx8&auto=format&fit=crop&w=1267&q=80")
```

### Objetivos de aprendizaje

En este tutorial vas a aprender como resolver problemas de iteración en R. En concreto, vas a aprender a diferenciar como resolver este tipo de problemas con:

* Bucles
* Funciones preconstruidas para situaciones especificas como `colMeans()` 
* Funciones echas a medida con `purrr`

Al final de este tutorial habrás practicado con varias herramientas de programación funcional en R.

## Repetición manual

Vamos a trabajar con unos nuevos datos de pinguinos llamados `pinguino_medidas` con variables, `pico`, `aletas` y `cabeza`. 

```{r pinguino-medidas, echo=TRUE}
pinguino_medidas
```

Imagina que quieres hacer la media de cada una de las tres variables de `pinguino_medidas`. Puedes calcular la media de cada una de las variables por separado en el espacio de abajo:

```{r meansep,exercise = TRUE}
___(pinguino_medidas$pico)
___(pinguino_medidas$aletas)
___(pinguino_medidas$cabeza)
```

```{r meansep-hint-1}
mean(pinguino_medidas$pico)
___(pinguino_medidas$aletas)
___(pinguino_medidas$cabeza)
```

```{r meansep-hint-2}
mean(pinguino_medidas$pico)
mean(pinguino_medidas$aletas)
___(pinguino_medidas$cabeza)
```

```{r mean-sep-hint-3}
mean(pinguino_medidas$pico)
mean(pinguino_medidas$aletas)
mean(pinguino_medidas$cabeza)
```

```{r meansep-solution}
mean(pinguino_medidas$pico)
mean(pinguino_medidas$aletas)
mean(pinguino_medidas$cabeza)
```

###

Puedes calcular las medias de esta manera, pero estarías repitiendo tu código (recuerda la norma de oro de no copiar y pegar más de dos veces una misma ejecución). Para evitar esto, podrías usar **un bucle _for_**.

## Iteración con bucles

Recueda las 3 cosas que tiene un bucle _for_:

* `output`: La parte del código que asigna suficiente espacio para guardar el resultado. 
* `secuencia`: La parte del código que determina cuantas veces hay que ejecutar el bucle.
* `cuerpo`: La parte del código que hace el trabajo.

Podrías construir un bucle para iterar sobre los `pinguino_medidas` y calcular la media de cada variable.

```{r, echo=TRUE}
salida_media <- vector("double", ncol(pinguino_medidas)) #output

for (i in seq_along(pinguino_medidas)) {     #sequencia
    salida_media[[i]] <- mean(pinguino_medidas[[i]])}     #cuerpo  
```

```{r, echo=TRUE}
salida_media
```

###

Imagina que en lugar de la media quieres calcular la mediana de cada variable, ¿cómo serías tu bucle _for_? ¡Inténtalo!

```{r ex1, exercise = TRUE}
salida_mediana <- vector("double", ncol(pinguino_medidas)) 

for (i in seq_along(pinguino_medidas)) {    
    salida_media[[i]] <- ___(pinguino_medidas[[i]])}      
salida_mediana
```

```{r ex1-solution}
salida_mediana <- vector("double", ncol(pinguino_medidas)) 

for (i in seq_along(pinguino_medidas)) {   
    salida_mediana[[i]] <- median(pinguino_medidas[[i]])} 
salida_mediana
```

###

A pesar de que los bucles __for__ permiten reducir el código, resultan más engorrosos porque hay que inicializar el output al comienzo, especificar el número de veces que el bucle se va a ejecutar, y escribir el bucle en sí. A continuación vamos a ver como podemos usar vectorización para simplificar este trabajo.


## Vectorización

En lugar de escribir un bucle __for__, podrías valerte del poder de [vectorizacion en R](https://swcarpentry.github.io/r-novice-gapminder-es/09-vectorization/) y calcular la media de todas las columnas con la function predefinida `colMeans()`:

```{r, echo=TRUE}
colMeans(pinguino_medidas)
```
Como ves `colMeans()`  ejecuta la media de todas las columnas del dataset, y de esta manera no es necesario escribir un bucle for. Esta función utiliza los principios de vectorización, lo que hace la escritura de código más concisa, fácil de leer y menos propensa a errores.

## Purrr

El paquete `purrr` ofrece otra alternativa más flexible que nos permite ejecutar iteraciones _sin tener que escribir bucles for_. Por ejemplo, puedes calcular la media de todas las columnas de `pinguino_medidas` así:

```{r, echo=TRUE}
map(pinguino_medidas,mean)
```

###

El diagrama de abajo te ayudará a recordar como functiona `map`. En el ejemplo de `pinguino_medidas` map aplica la functión `mean()` a todas las variables de pinguino_medidas.

```{r echo=FALSE, out.width = "70%"}
knitr::include_graphics("https://raw.githubusercontent.com/rstudio-education/primers/master/iterate/map/www/images/map.png")
```

También puedes utilizar map con el pipe:

```{r, echo=TRUE}
pinguino_medidas %>% map(mean)
```

###

¿Qué tendrías que poner para calcular la mediana de todas las variables de `pinguino_medidas`?

```{r ex2, exercise = TRUE}
pinguino_medidas %>% map(___)

```

```{r ex2-solution}
pinguino_medidas %>% map(median)
```

```{r ex2-check}
grade_code("En este caso la media y la mediana tienen el mismo valor.")
```

###

También puedes pasar tu propia función a map. Por ejemplo, puedes crear una función que calcule la media después de quitar NAs. 

```{r, echo= TRUE}
mi_media_sin_nas <- function(x) {
  x %>%
    mean(na.rm = TRUE)}
```

Y calcula la media para variables numéricas de pinguinos con tu nueva función.

```{r, echo=TRUE}
pinguinos %>% 
  select_if(is.numeric) %>%
  map(mi_media_sin_nas)
```

### Las funciones map

Si has estado atenta, habrás visto que map retorna por defecto una lista. Hay otras functiones de map que devuelven distintos tipos de output.

Función map    | Output                                               | Ejemplo
----------------- | ------------------------------------------------- | -------
**map**             | Lista                         | `map(pinguino_medidas, mean)`
**map_dbl**             | Vector de dobles                  | `map_dbl(pinguino_medidas, mean)`
**map_chr**             | Vector de caracteres                         | `map_chr(pinguino_medidas, mean)`
**map_int**             | Vector de enteros                         | `map_int(pinguino_medidas, length)`
**map_lgl**             | Vector lógico                         | `map_lgl(pinguino_medidas, is.numeric)`


¿Cómo podrías retornar el resultado anterior como un vector de dobles?

```{r ex3, exercise = TRUE}
pinguinos %>% select_if(is.numeric) %>% ___(mi_media_sin_nas)
```

```{r ex3-hint-1}
"Mira la tabla anterior para ver cual de las funciones map devuelve un vector de dobles."
```

```{r ex3-solution}
pinguinos %>% select_if(is.numeric) %>% map_dbl(mi_media_sin_nas)
```

## En resumen

Enhorabuena. Has practicado iteración con bucles for y con map().

map() tiene más cosas útiles como atajos de sintaxis. Puedes practicar como utilizar atajos de sintaxis en la seccion Extra de este tutorial.

## Extra

`pinguino_lista_Torgersen` es una lista con las medidas de pinguinos de la isla Torgersen. Puedes inspeccionar esta lista en el ejercicio de abajo. 

```{r pinguino_lista-1, exercise = TRUE}
pinguino_lista_Torgersen
```

```{r pinguino_lista-solution}
pinguino_lista_Torgersen
```

###  

¿Puedes calcular el incremento en `masa_corporal_g` desde 2017 hasta 2019 en la lista pinguino_lista_Torgersen? Esta puede parecer una tarea fácil, y va a ser la base de nuestro ejercicio de iteracion. 

### Manejo de datos 

En primer lugar tienes que acceder a los valores de `masa_corporal_g`. Trabajar con estos valores puede resultar difícil porque estan almacenados dentro de la lista `pinguino_lista_Torgersen`. 

Puedes extraer los valores con la función `pluck()` del paquete  purrr's.

### pluck() 

`pluck()` extrae un elemento de una lista, por nombre o por posición. Devuelve el contenido del elemento tal como está, sin rodearlo en una nueva lista.

Por ejemplo, la llamada `pluck()` extrae los valores de `masa_corporal_g` de `pinguino_lista_Torgersen`.


```{r pluck, exercise = TRUE}

```

```{r pluck-hint-1}
"El nombre del elemento que deseas extraer es masa_corporal_g"
```

```{r pluck-hint-2}
"No te olvides de poner comillas alrededor de masa_corporal_g cuando lo pases a pluck()."
```

```{r pluck-solution}
pinguino_lista_Torgersen %>% pluck("masa_corporal_g")
```

```{r pluck-check}
grade_code("Muy bien. pluck() hace lo mismo que pinguino_lista_Torgersen$masa_corporal_g, pero pluck() es más fácil
de leer (y más fácil de pasar a map()). `pluck()` devuelve los valores de `masa_corporal_g` como vector, lo que nos facilita trabajar con vectorización.")
```

### Lista de listas

`pinguino_lista` es una lista con las medidas de pinguinos de las islas Torgersen, Biscoe y Dream. Selecciona Enviar respuesta para ver el contenido de `pinguino_lista`.

```{r pinguino-lista, exercise = TRUE}
pinguino_lista
```

```{r pinguino-lista-solution}
pinguino_lista
```

```{r pinguino-lista-check}
grade_code("Cada elemento de pinguino_lista tiene la misma estructura que pinguino_lista_Torgersen.")
```

### 

Podemos extraer la masa_corporal_g de los pinguinos de cada una de las listas de las 3 islas con la combinación map and pluck:

```{r map-pluck, exercise = TRUE}
pinguino_lista %>% ___(___, "masa_corporal_g")
```

```{r  map-pluck-hint}
"Acuerdate que map() toma un vector sobre el que iterar, una función para aplicar a cada elemento del vector y luego cualquier argumento que la función necesite."
```

```{r map-pluck-solution}
pinguino_lista %>% map(pluck, "masa_corporal_g")
```

### Atajos

Puede que todavia no sepas como hacer um modelo lineal con R. Por el momento toma mi palabra que el código de abajo ajusta un modelo lineal para calcular la masa corporal en funcion del año para la lista `pinguino_lista_Torgersen`.  

```{r lm, exercise = TRUE}
lm(masa_corporal_g ~ anio, data = pinguino_lista_Torgersen)
```

```{r lm-hint}
"Deja el código como está y haz clic en Enviar respuesta."
```

```{r lm-solution}
lm(masa_corporal_g ~ anio, data = pinguino_lista_Torgersen)
```

```{r lm-check}
grade_code("La función lm (linear model en ingles) construye un modelo lineal  ")
```

###

Imagina ahora que quieres hacer este modelo para cada una de las 3 islas incluidas en `pinguino_lista`. Esto no va a ser muy difícil
 porque ya hemos hecho lo más complicado que es escribir el modelo. Sin embargo hay 3 reglas que debemos saber. 
 
### Regla 1 - ~

En primer lugar, coloca  `~` al comienzo de la expresión. Esto alerta a `map()` que le estás dando una expresión para que se ejecute:

Por ejemplo, `~lm(masa_corporal_g ~ anio, data = pinguino_lista)`

### Regla 2 - .x

En segundo lugar, reemplaza el _nombre de la lista para manipular_ por `.x` donde la lista aparece en la función.

Por ejemplo,`pinguino_lista %>% ~lm(masa_corporal_g ~ anio, data = .)`

Esto le dice a `map()` dónde usar el elemento dentro de su expresión. Si una expresión usa cada elemento varias veces, deberás insertar varios `.x`.

### Regla 3 - map

En tercer lugar, pasa la funcion a map().

Por ejemplo, `pinguino_lista %>% map(~lm(masa_corporal_g ~ anio, data = .)`

###

Imagina que quieres ejecutar un modelo lineal para calcular la masa corporal en funcion del año y el tamano de la cabeza. Te muestro el ejemplo de como seria este modelo para la lista `pinguino_lista_Torgersen`. ¿Podrías implementar este modelo para los datos `pinguino_lista`? 

```{r exp_lm, exercise = TRUE}
lm(masa_corporal_g ~ anio + cabeza, data = pinguino_lista_Torgersen)

```

```{r exp_lm-hint-1}
"Recuerda las tres reglas para las expresiones map(): 1) empieza con a ~, 2) haz referencia a la lista con .x y 3) Pasa la función a mapa ()."
```

```{r exp_lm-solution}
pinguino_lista %>%
  map(~ lm(masa_corporal_g ~ anio + cabeza, data = .x))
```

```{r exp_lm-check}
grade_code("Buen trabajo. Has completado as sección extra de los ejercicios de iteración.")
```

