---
title: "Iteraction"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Use iteration to reduce duplication in your code.
---

<!-- Add JavaScript code for making the exercise code larger -->
<script language="JavaScript" src="js/exercise-font-size.js"></script>

```{r setup, include=FALSE}
# load packages-----------------------------------------------------------------

library(learnr)
library(tidyverse)
library(palmerpenguins)
library(gradethis)
knitr::opts_chunk$set(echo = FALSE, exercise.checker = gradethis::grade_learnr)

# set options for exercises and checking ---------------------------------------
penguins_measurements <- tibble(
  bill = c(100,200,300),
  fins= c(200,400,600),
  head = c(400,800,1200))

penguins_list_Torgersen <- list(
  body_mass_g = c(4433, 4338, 4559),
  year= c(2017,2018,2019),
  head = c(450,850,1250))

penguins_list_Biscoe <- list(
  body_mass_g = c(4868, 3504, 5219),
  year= c(2017,2018,2019),
  head = c(903, 1125,  514))

penguins_list_Dream <- list(
  body_mass_g = c(4280, 5134, 5183),
  year= c(2017,2018,2019),
  head = c(588,  718, 1094))


penguins_list <- list(
  penguins_list_Torgersen,
  penguins_list_Biscoe,
  penguins_list_Dream
)

names(penguins_list)<-c("Torgersen","Biscoe","Dream")


my_mean_wo_nas <- function(x) {
  x %>%
    mean(na.rm = TRUE)}

# hide non-exercise code chunks ------------------------------------------------
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(echo = FALSE)
```

## Introduction

```{r imagen-peng, fig.margin = TRUE, echo = FALSE, out.width = "100%", fig.cap="Unsplash: Andrea Gradilone"}
knitr::include_graphics("https://images.unsplash.com/photo-1506842605307-bbcbd9ea121f?ixlib=rb-1.2.1&ixid=MnwxMjA3fDB8MHxwaG90by1wYWdlfHx8fGVufDB8fHx8&auto=format&fit=crop&w=1267&q=80")
```

### Learning objectives

In this tutorial you will learn how to solve iteration problems in R. Specifically, you will learn to solve these types of problems with:

* For loops
* Pre-built functions for specific situations like `colMeans()`
* Tailored functions of `purrr` package

By the end of this tutorial you will have practice using tools to do functional programming in R.

## Manual repetition

We are going to work with some a penguin dataset called `penguin_measures` with variables,`bill`, `fins` and ` head`. Just clic on Run Code to see this dataset.

```{r pinguino_medidas1, exercise = TRUE}
penguins_measurements
```

###

Suppose you want to calculate the average each of the three variables in `penguins_measurements`. You could calculate the average of each variable individually by calling the appropriate function on the box below:

```{r meansep,exercise = TRUE}
___(penguins_measurements$bill)
___(penguins_measurements$fins)
___(penguins_measurements$head)
```

```{r meansep-hint-1}
mean(penguins_measurements$bill)
___(penguins_measurements$fins)
___(penguins_measurements$head)
```

```{r meansep-hint-2}
mean(penguins_measurements$bill)
mean(penguins_measurements$fins)
___(penguins_measurements$head)
```

```{r mean-sep-hint-3}
mean(penguins_measurements$bill)
mean(penguins_measurements$fins)
mean(penguins_measurements$head)
```

```{r meansep-solution}
mean(penguins_measurements$bill)
mean(penguins_measurements$fins)
mean(penguins_measurements$head)
```

```{r meansep-check}
grade_code("You can indeed calculate the mean of each variable in this way, but you would be repeating your code (remember the rule of thumb: never copy and paste your code more than twice). To avoid this, you could use  **a for loop**.")
```

###

## Iteraction with for loops

Remember the three components of a for loop:

* `output`: The part of the code that allocates enough space to save the resulto. 
* `sequence`: The part of the code that determines how many times to execute the loop.
* `body`: The part of the code that does the job.

You could build a for loop to iterate over the `penguins_measurements` and calculate the mean of each variable, like so:

```{r, echo=TRUE}
output_mean <- vector("double", ncol(penguins_measurements)) 

for (i in seq_along(penguins_measurements)) {    
    output_mean[[i]] <- mean(penguins_measurements[[i]])}      
```

```{r, echo=TRUE}
output_mean
```

###

Imagine that instead of the mean you want to calculate the median of each variable, how would your `for loop` look like? Give it a try!

```{r ex1, exercise = TRUE}
output_median <- vector("double", ncol(penguins_measurements)) 

for (i in seq_along(penguins_measurements)) {    
    output_median[[i]] <- ___(penguins_measurements[[i]])}      
output_median
```

```{r ex1-solution}
output_median <- vector("double", ncol(penguins_measurements)) 

for (i in seq_along(penguins_measurements)) {    
    output_median[[i]] <- median(penguins_measurements[[i]])}      
output_median
```

###

Although for loops allow reducing repetition on our code, they are more cumbersome because we have to 1) initialize the output at the beginning, 2) specify the number of times that the loop is to be executed, and 3) write the loop itself. Next we will see how we can use vectorization to simplify this task.

###

### Vectorization

Instead of writing a for loop, we can make use [the power of vectorization in R](https://swcarpentry.github.io/r-novice-gapminder/09-vectorization/) and calculate the mean of all columns with the predefined function `colMeans()`:

```{r, echo=TRUE}
colMeans(penguins_measurements)
```
As you see, `colMeans()` executes the average of all the columns in the dataset, and in this way we do not need to write a for loop. This function uses the principles of vectorization, which makes code writing more concise, easier to read, and less error prone.

### Purrr

The `purrr` package offers another more flexible alternative that allows executing iterations _without having to write for_ loops. For example, we can calculate the mean of all the columns of `penguin_measures` like this:

```{r, echo=TRUE}
map(penguins_measurements,mean)
```

###
The diagram below will help you remember how map works.

```{r echo=FALSE, out.width = "100%"}
knitr::include_graphics("https://raw.githubusercontent.com/rstudio-education/primers/master/iterate/map/www/images/map.png")
```


In the `penguins_measurements` example map applies the `mean()` function to all variables.

###

You can also use map with the pipe: 

```{r, echo=TRUE}
penguins_measurements %>% map(mean)
```

###

How would you calculate the median of each `penguins_measurements` column?

```{r ex2, exercise = TRUE}
penguins_measurements %>% ___(___)
```

```{r ex2-hint-1}
penguins_measurements %>% map(___)
```

```{r ex2-solution}
penguins_measurements %>% map(median)
```

```{r ex2-check}
grade_code("Excellent. In this case the mean and median have the same value.")
```

###

You can also pass **your own function** to map. For example, you can create a function that calculates the mean after removing NAs.

```{r, echo= TRUE}
my_mean_wo_nas <- function(x) {
  x %>%
    mean(na.rm = TRUE)}
```

###

And calculate the mean for the numerical variables of penguins using your new function.

```{r, echo=TRUE}
penguins %>% 
  select_if(is.numeric) %>%
  map(my_mean_wo_nas)
```

###

### Map functions

You may have notice that the default output of map is a list. There are similar functions to map that return different types of output.

Map function    | Output                                               | Example
----------------- | ------------------------------------------------- | -------
**map**             | List                        | `map(penguins_measurements, mean)`
**map_dbl**             | Double vector                 | `map_dbl(penguins_measurements, mean)`
**map_chr**             | Character vector                        | `map_chr(penguins_measurements, mean)`
**map_int**             |  Integer vector                         | `map_int(penguins_measurements, length)`
**map_lgl**             | Logical vector                         | `map_lgl(penguins_measurements, is.numeric)`

###

How you could get the output of the previous operation as double vector? 

```{r ex3, exercise = TRUE}
penguins %>%
  select_if(is.numeric) %>%
  ___(my_mean_wo_nas)
```

```{r ex3-hint-1}
"Check on the previous table which map function returns double vector."
```

```{r ex3-solution}
penguins %>% 
  select_if(is.numeric) %>%
  map_dbl(my_mean_wo_nas)
```

## Wrap up

Well done! You've practice iteration with for loops and with map(), but of course there is more to learn about map()!

Go ahead and explore the Extra section of this tutorial to learn more about map() and puuur.

## Extra

`penguins_list_Torgersen` is a list that contains penguin measurements from Torgersen Island. You can inspect this list in the exercise below:

```{r pinguino_lista-1, exercise = TRUE}
penguins_list_Torgersen
```

###  

Could you calculate the increase in `body_mass_g` from 2017 to 2019 in the list pinguino_list_Torgersen? This may seem like an easy task, and it will be the basis of our iteration exercise.

### Data wrangling

First you would have to get `body_mass_g` values. This can be a bit tricky since the values are stored within the `penguins_list_Torgersen` list. You can extract the values using the `pluck()` function from the purrr's package.

### pluck() 

`pluck()` extracts items from a list, by name or by position and it returns the content of the element as is, without surrounding it in a new list.

For example, you could use `pluck()` to extract `body_mass_g` from `penguins_list_Torgersen`.

```{r pluck, exercise = TRUE}

```

```{r pluck-hint-1}
"The name of the element you want to extract is body_mass_g"
```

```{r pluck-hint-2}
"Don't forget to put quotes around body_mass_g when passing it to pluck ()."
```

```{r pluck-solution}
penguins_list_Torgersen %>% pluck("body_mass_g")
```

```{r pluck-check}
grade_code("pluck() does the same as penguins_list_Torgersen$body_mass_g, but pluck() is easier
to read (and easier to pass to map ()). `pluck()` returns `body_mass_g` values as a vector, which makes it easier for us to work with vectorization.")
```

### List of list 

`penguins_list` is a list of penguin measurements from Torgersen, Biscoe and Dream Islands. Select Send reply to view the content of `penguins_list`.

```{r pinguino-lista, exercise = TRUE}
penguins_list
```

```{r pinguino-lista-solution}
penguins_list
```

```{r pinguino-lista-check}
grade_code("Each element of penguins_list has the same structure as penguins_list_Torgersen")
```

### 

We can extract body_mass_g from penguins from each of the lists of the 3 islands with the combination map and pluck:

```{r map-pluck, exercise = TRUE}
penguins_list %>% ___(___, "body_mass_g")
```

```{r  map-pluck-hint}
"Remember that map() takes a vector to iterate over, a function to apply to each element of the vector and then whatever argument the function needs."
```

```{r map-pluck-solution}
penguins_list %>% map(pluck, "body_mass_g")
```

### Purrr style

You may not yet know how to make a linear model with R. For now, take my word for it that the code below fits a linear model to calculate body mass as a function of year for the list `penguin_list_Torgersen`.  

```{r lm, exercise = TRUE}
lm(body_mass_g ~ year, data = penguins_list_Torgersen)
```

```{r lm-hint}
"Leave the code as is and click Submit Answer."
```

```{r lm-solution}
lm(body_mass_g ~ year, data = penguins_list_Torgersen)
```

```{r lm-check}
grade_code("lm function is used to build a linear model")
```

###

Now imagine that you want to make this model for each of the 3 islands included in `penguin_list`. This is not going to be very difficult
 because we have already done the most complicated thing, which is writing the model. However there are 3 rules of purrr that we must know to complete this task.
 
### Rule 1 - ~

First, we need to put `~` at the beginning of the expression. This alerts `map()` that we are giving an expression to execute:

For example, `~lm(body_mass_g ~ year, data = penguins_list)`

### Rule 2 - .x

Second, we need to replace the _name of the list to manipulate_ with `.x`

For example,`penguins_list %>% ~lm(body_mass_g ~ year, data = .)`

This tells `map()` where to use the element within the expression. If an expression uses each element multiple times, we will need to insert multiple `.x`s into the expression.

### Rule 3 - map

Third, we need to pass the function to map().

For example, `penguins_list %>% map(~lm(body_mass_g ~ year, data = .)`

###

Suppose you want to run a linear model to calculate body mass based on year and head size. I show below how this model would be for the list `pinguino_lista_Torgersen`. Could you implement this model for the `penguin_list` data?

```{r exp_lm, exercise = TRUE}
lm(body_mass_g ~ year + head, data = penguins_list_Torgersen)

```

```{r exp_lm-hint-1}
"Remember the three rules for map() expressions: 1) start with a ~, 2) refer to the list with .x and 3) Pass the function to map()."
```

```{r exp_lm-solution}
penguins_list %>%
  map(~ lm(body_mass_g ~ year + head, data = .x))
```

```{r exp_lm-check}
grade_code("You have now completed the extra section of the iteration exercises.")
```

